<a href="https://colab.research.google.com/github/wenshanliu-slu0/hds5210-2023/blob/main/week06/week06_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In this GitHub repository, there is a file called `allergies.json` that contains a list of patient allergies.  You will need to download this [file from here](https://raw.githubusercontent.com/paulboal/hds5210-2023/main/week06/allergies.json) and then upload it into Google Colab to run these examples. It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE="allergies.json"
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here

def allergy_count(ALLERGIES_FILE):

  """(json_file)-->int

    This function is to calcuate the numbers of the allergy list.

    This version program will calcuate the numbers of the allergy list from the improrting json file, and return 'int' datatpye at the last.

Example:
>>> allergy_count(ALLERGIES_FILE)
4

  """

  with open(ALLERGIES_FILE,'r') as file:
     data = json.load(file)
     if 'entry' in data and isinstance(data['entry'], list):
                return len(data['entry'])
     else:
                return 0  # Return 0 if 'entry' is missing or not a list
```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

```python
import doctest
doctest.run_docstring_examples(allergy_count, globals(), verbose=True)
```

Output:
```
PYDEV DEBUGGER WARNING:
sys.settrace() should not be used when the debugger is being used.
This may cause the debugger to stop working correctly.
If this is needed, please check: 
http://pydev.blogspot.com/2007/06/why-cant-pydev-debugger-work-with.html
to see how to restore the debug tracing back correctly.
Call Location:
  File "/usr/lib/python3.10/doctest.py", line 1501, in run
    sys.settrace(save_trace)

Finding tests in NoName
Trying:
    allergy_count(ALLERGIES_FILE)
Expecting:
    4
ok
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
# Put your solution here
def patient_count(ALLERGIES_FILE):

  """(json_file)-->int

    This function is to find the number of paitents.

    This version program will find the numbers of paitents from the improrting json file, and return 'int' datatpye at the last.


Example:
>>> patient_count(ALLERGIES_FILE)
2

  """

  patient_num_count =set()  #initialize the parameter

  with open(ALLERGIES_FILE,'r') as file: #open and load the JSON
      data = json.load(file)

  if 'entry' in data:  # and (data['entry'], list): # IF 1) 'entry' in data and 2) entry is a list
        for entry in data['entry']:

          if 'resource' in entry and 'patient' in entry['resource']:
            patient = entry['resource']['patient']['display']
            patient_num_count.add(patient)

        return len(patient_num_count)

```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

```python
import doctest
doctest.run_docstring_examples(patient_count, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    patient_count(ALLERGIES_FILE)
Expecting:
    2
ok
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
# Put your solution here

def allergy_per_patient(ALLERGIES_FILE):
    """(json_file)-->dict

    This function is to list how many allergies each patient has.

    This version program will how many allergies each patient has from improrting json file , and return 'dictionary' datatpye at the last.
    For instance  : {'Jason Argonaut': 3, 'Paul Boal': 1}


Example:
>>> allergy_per_patient(ALLERGIES_FILE)
{'Jason Argonaut': 3, 'Paul Boal': 1}

  """


    # Initialize a dictionary to store the count of allergies per patient
    allergy_count = {}

    with open(ALLERGIES_FILE,'r') as file: #open and load the JSON
     data = json.load(file)


    for entry in data["entry"]:
        patient_name = entry["resource"]["patient"]["display"]

        # Check if the patient is already in the dictionary, if not, add them
        if patient_name not in allergy_count:
            allergy_count[patient_name] = 0


        # allergy_count add +1 if paient in the dict
        allergy_count[patient_name] = allergy_count[patient_name]+1

    return allergy_count
```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

```python
import doctest
doctest.run_docstring_examples(allergy_per_patient, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_per_patient(ALLERGIES_FILE)
Expecting:
    {'Jason Argonaut': 3, 'Paul Boal': 1}
ok
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_list(ALLERGIES_FILE):

  """(json_file)-->list

    This function is to list the paients name, allergy, and reaction.

    This version program will read the improrting json file and and return 'list' datatpye at the end.
    The list should like this :['Name', 'allergy', 'reaction'].
    For instance:  [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
    ['Paul Boal', 'PENICILLIN G', 'Bruising'],
    ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
    ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]

Example:
>>> allergy_list(ALLERGIES_FILE)
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],['Paul Boal', 'PENICILLIN G', 'Bruising']]

  """




  total_list =[] #Initialize a list

  with open(ALLERGIES_FILE,'r') as file: #open and load the JSON
     data = json.load(file)

     for entry in data["entry"]:
        patient_name = entry["resource"]["patient"]["display"]
        allergy_text =entry["resource"]["substance"]["text"]
       # allergy_reaction= entry["resource"]["reaction"]["text"]

        if 'reaction' in entry['resource'] and len(entry['resource']['reaction']) > 0:
            reaction = entry['resource']['reaction'][0]['certainty']
            if 'manifestation' in entry['resource']['reaction'][0] and len(entry['resource']['reaction'][0]['manifestation']) > 0:
                manifestation = entry['resource']['reaction'][0]['manifestation'][0]['text']
            else:
                manifestation = None
        else:
            reaction = None
            manifestation = None

        # Add the extracted information to the output list
        total_list.append([patient_name, allergy_text, manifestation])



  return total_list
```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_list(ALLERGIES_FILE, patient):

    """(json_file)-->list

    This function is to list the paients name, allergy, and reaction.

    This version program will list each paient's name, allergy, and reaction has from improrting json file.

Example:
>>> allergy_list(ALLERGIES_FILE)
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],['Paul Boal', 'PENICILLIN G', 'Bruising']]

  """

    # Read the JSON file
    with open(ALLERGIES_FILE, 'r') as file:
        data = json.load(file)

    #Initialize a new list to show paient's allegy
    allergy_entries = data.get('entry', [])

    # Search for allergies for the specified patient
    allergies = []
    for entry in allergy_entries:
        resource = entry.get('resource', {})
        patient_info = resource.get('patient', {}).get('display', '')
        if patient_info == patient:
            substance = resource.get('substance', {}).get('text', '')
            allergies.append(substance)

    return allergies

def allergy_reaction(ALLERGIES_FILE, patient, substance):

    """(json_file,string, string)-->string

    This function is to find the reaction when the paient got the specific substance.

    This version program is to find the reaction when the paient got the specific substance from improrting json file.
    It would return 'string' datatype at the end.


Example:
>>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
'Hives'
  """


    # Get the list of allergies for the specified patient
    allergies = allergy_list(ALLERGIES_FILE, patient)

    # Check if the specified substance is in the patient's allergy list
    if substance in allergies:
        # Find the reaction for the specified substance
        with open(ALLERGIES_FILE, 'r') as file:
            data = json.load(file)

        allergy_entries = data.get('entry', [])
        for entry in allergy_entries:
            resource = entry.get('resource', {})
            patient_info = resource.get('patient', {}).get('display', '')

            if patient_info == patient and resource.get('substance', {}).get('text', '') == substance:
                reactions = resource.get('reaction', [])
                if reactions:
                    return reactions[0].get('manifestation', [])[0].get('text', None)

    else:
     return None





```

```python
allergy_list(ALLERGIES_FILE ,'Paul Boal')
```

Output:
```
['PENICILLIN G']
```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

```python
import doctest
doctest.run_docstring_examples(allergy_reaction, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
Expecting:
    'Hives'
ok
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

